```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04A.ipynb")
```

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

```

### Pre - Homework 04A: Boolean and Count

Please complete the questions below before starting on the homework. 


#### Boolean
![Boolean Expressions](https://raw.githubusercontent.com/lbcc-data-science/DataScienceForAll_S26_LD/refs/heads/main/hw/hw04A/BooleanExpressions.png)

**Question 1.** Use a boolean operator to determine if 4 is less than 9.

```python
# put expression here 
...
```

**Question 2.** Use a boolean operator to determine if 3 is equal to 'three'.

```python
# put expression here 
...
```

##### Read in the `movies.csv`file below. 

```python
# Just run this cell
# Read in the movies csv file
movies = Table().read_table("movies.csv")
movies.show()
```

**Question 3.** Set `profit` equal to an array of boolean expressions that tells us which movies made a profit. Remember that profit comes from making more money (revenue) than the money spent (budget). 

```python
 # We can compare arrays if they're the same size 
# write the expression after profit
profit = ...
profit 
```

**Question 4.** Assign `profit_movies_amount` to the amount of movies that made a profit. Use the `np.count_nonzero` function on the `profit` array.
In Boolean, `True` = 1 and `False` = 0. Using `np.count_nonzero` on boolean results counts the `True` responses.

```python
profit_movies_amount = ...
profit_movies_amount
```

**Question 5.** Assign `profit_movies` to a table with the following columns: `Movie`, `Revenue (Millions)`, and `Profit?`, where `Profit?` is the column with the `True/False` values found in the `profit` array. 

```python
profit_movies = ...
profit_movies
```

# Homework 04A: Applying Functions and Iteration

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests. Each time you start your server, you will need to execute this cell again to load the tests.

**Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells.**

## 1. 22-23 LBCC Women's Basketball Season


Let's take a look at how well the Long Beach City College (LBCC) Women's Basketball team performed in the 21-22 season. A basketball game is divided into four periods, called quarters. The number of points LBCC scored in each quarter, and the number of points their opponent scored in each quarter are stored in a table called `lbcc_bb.csv`.

```python
# Just run this cell
# Read in the lbcc_bb csv file
games = Table().read_table("lbcc_bb.csv")
games.show(5)
```

Let's start by finding the total points each team scored in a game.

**Question 1.** Write a function called `game_totals`.  It should take four arguments, where each argument is the team's score for that quarter. It should return the team's total score for that game.


<!--
BEGIN QUESTION
name: q1_1
manual: false
-->

```python
def game_totals(...):
    '''Returns the total score calculated by adding up the score of each quarter'''
    ...


game_totals(14, 7, 3, 0) #DO NOT CHANGE THIS LINE
```

```python
grader.check("q1_1")
```

**Question 2.** Create a new table `results` with three columns in this *specific* order: `Opponent Name`, `LBCC Score`, `Opponent Score`. You will have to create the `LBCC Score` and `Opponent Score` columns. Use the function `game_totals` you just defined in the previous question for this problem.

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating final_scores, check that `LBCC_scores` and `opponent_scores` output what you want. 


<!--
BEGIN QUESTION
name: q1_2
manual: false
-->

```python
LBCC_scores = ...
opponent_scores = ...
results = ...
results

```

```python
grader.check("q1_2")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `(n+1)`th row of a data from the table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
games.row(10)
```

```python
# Just run this cell
games.row(10).item("LBCC 4Q")
```

**Question 3.** We want to see for a particular game whether or not LBCC won. Write a function called `did_LBCC_win`.  It should take one argument: a row object from the `results` table. It should return either `True` if LBCC's score was greater than the Opponent's score, and `False` otherwise.


<!--
BEGIN QUESTION
name: q1_3
manual: false
-->

```python
def did_LBCC_win(row):
...
   
```

```python
grader.check("q1_3")
```

**Question 4.** Did LBCC win the majority of the games in the 22-23 season? Using the `results` table, assign `wins` to an array of `True` and `False` values that correspond to whether or not LBCC won. Add the `wins` array to the `results` table as a column called `Win?` and assign this to a new table called `results_with_wins`. Then, respectively assign the number of wins and losses LBCC had to `LBCC_wins` and `LBCC_losses`, but be sure that this is done completely in Python.

*Hint*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`. 


<!--
BEGIN QUESTION
name: q1_4
manual: false
-->

```python
wins = ...
results_with_wins = ...
LBCC_wins = ...
LBCC_losses = ...



# Don't delete or edit the following line:
print(f"In the 22-23 Season, the LBCC Women's Basketball team won {LBCC_wins} games and lost {LBCC_losses} games.")
```

<!-- END QUESTION -->

## Congratulations! You're done with Homework 04A!
**Complete the following steps to prepare your work for submission.**


1. Click on **File** in the menu bar, and choose **Save Notebook**.
2. Then click on **File** again, and choose **Save and Export Notebook As** and select PDF. (Be patient. This may take a while.) 
3. Then click on **File** again, choose **Download**.
4. Find the two files that you have just downloaded (one should end with *.ipynb* and the other with *.pdf*).
5. Rename each file in the following format, with **no spaces**: *LastName FirstInitial AssignmentType TwoDigitAssignmentNumber* 


For example, Taylor Swift's homework 04A will be submitted in Canvas as: SwiftTHW04A.ipynb and SwiftTHW04A.pdf 




```python
grader.check("q1_4")
```

Tests for q1_1 (points: None)
```
>>> game_totals(2, 3, 6, 1)
12
```
```
>>> game_totals(-2,3,5,-10)
-4
```


Tests for q1_2 (points: None)
```
>>> results.num_columns
3
```
```
>>> set(['Opponent Name', 'LBCC Score', 'Opponent Score']) == set(results.labels)
True
```


Tests for q1_3 (points: None)
```
>>> 'did_LBCC_win' in globals()
True
```
```
>>> did_LBCC_win(results.row(1))
True
```


Tests for q1_4 (points: None)
```
>>> 0 <= LBCC_wins <= 15
True
```
```
>>> 0 <= LBCC_losses <= 14
True
```
