```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, run the previous cell to load the provided tests.

**Recommended Readings:**

- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more hidden tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

## 1. Scary Arithmetic

An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

**Question 1** Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day? Read the response below and then assign either 1, 2 or 3 to the variable `adt_ad` below.

1. Yes, the data clearly show that burglaries are significantly more frequent between Memorial Day and Labor Day compared to other times of the year.
2. No.  Memorial day is the last Monday of May and Labor day is the first Monday of September, which is a little over 3 months or 25% of the year.  If we expect burglaries to be uniformly distributed throughout the year, then 25% of burglaries occurring in about 25% of the year makes sense. 
3. No, the data only state that 25% of burglaries occur during this period, but without information on burglary rates for the rest of the year, we cannot determine if this period is more dangerous




```python
adt_ad = ...
```

```python
grader.check("q1")
```

## 2. Characters in Little Women


In lecture, we counted the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://inferentialthinking.com/chapters/01/3/1/Literary_Characters.html?highlight=little%20women). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# Just run this cell.

# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

**Question 1.** Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below.

1. 250
2. 390
3. 440
4. 32,000
5. 40,000


**Note:** If you run into a `NameError: name 'grader' is not defined` error in the autograder cell below (and in any assignment), please re-run the first cell at the very top of this notebook!


```python
characters_q1 = ...
```

```python
grader.check("q2_1")
```

The test above checks that your answers are in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.

**Question 2.** Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below.

1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods


```python
characters_q2 = ...
```

```python
grader.check("q2_2")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, check out [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) in the textbook.

## 3. Names and Assignment Statements

**Question 1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer. 

1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q1 = ...
```

```python
grader.check("q3_1")
```

**Question 2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. The name `plus` isn't a built-in operator; instead, addition uses `+`.


```python
names_q2 = ...
```

```python
grader.check("q3_2")
```

**Question 3.** Run the following cell.

```python
x = 2
y = 3 * x
x = 4
```

What is `y` after running this cell, and why? Choose the best explanation and assign 1, 2, 3, or 4 to `names_q3` below to indicate your answer. 

1. `y` is equal to 6, because the second `x = 4` has no effect since `x` was already defined.

2. `y` is equal to 6, because `x` was 2 when `y` was assigned, and 3 * 2 is 6.

3. `y` is equal to 12, because `x` is 4 and 3 * 4 is 12.

4. `y` is equal to 12, because assigning `x` to 4 will update `y` to 12 since `y` was defined in terms of `x`.


```python
names_q3 = ...
```

```python
grader.check("q3_3")
```

## 4. Differences Between Majors

Berkeley’s Office of Planning and Analysis provides data on numerous aspects of the campus. Adapted from the OPA website, the table below displays the number of degree recipients in three majors in the 2008-2009 and 2017-2018 academic years.

| Major                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Gender and Women's Studies         |      17      |    28       |
| Linguistics                        |      49      |    67       |
| Rhetoric                           |      113     |    56       |



**Question 1.** Suppose you want to find the **biggest** absolute difference between the number of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 49 from 67) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers. 

```python
biggest_change = ...
biggest_change
```

```python
grader.check("q4_1")
```

**Question 2.** Which of the three majors had the **smallest** absolute difference? Set `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major. 


```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q4_2")
```

**Question 3.**  For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2009}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors.

```python
gws_relative_change = (abs(...) / 17) * 100
linguistics_relative_change = ...
rhetoric_relative_change = ...
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

```python
grader.check("q4_3")
```

**Question 4.** Set `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the biggest relative change. 


```python
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q4_4")
```

## 5. Nearsightedness Study

[Myopia](https://en.wikipedia.org/wiki/Myopia), or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

**Question 1.** The data were gathered by the following procedure, reported in the study. "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child's light exposure both at present and before the age of 2 years." Was this study observational, or was it a controlled experiment? Assign either 1 or 2 to the variable `study_type` below.
1. Observational study – The researchers collected data based on parents’ reports without manipulating light exposure conditions.
2. Controlled experiment – The researchers assigned different lighting conditions to children and observed the effects on myopia.


```python
study_type = ...
```

```python
grader.check("q5_1")
```

**Question 2.** The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded the following: "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."
Do the data support this statement? Assign 1, 2, 3 or 4 to the variable `associated` below.

1. No, the study does not provide enough data to determine any relationship between ambient light exposure and myopia.
2. No, the data only show correlation, but they do not establish that light exposure causes myopia. Other factors may be responsible for the observed differences.
3. Yes, the data show a clear association between increased ambient light exposure during sleep and higher prevalence of myopia in children. 
4. Yes, since children exposed to more light had a higher percentage of myopia, this proves that light exposure causes myopia.

```python
associated = ...
```

```python
grader.check("q5_2")
```

**Question 3.** On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the conclusion of the study claim that night light causes nearsightedness? *Hint: Look back at the quote in question 2 to see what language the study used in the conclusion.*  Assign 1, 2 or 3 to the variable `causal_relationship` below.

1. Yes, the study’s wording strongly suggests that night light exposure is the primary cause of myopia.
2. Yes, the study concluded that night light exposure directly causes nearsightedness in children.
3. No, the study only found an association between night light use and myopia, not a causal relationship.


```python
causal_relationship = ...
```

```python
grader.check("q5_3")
```

**Question 4.** The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.
Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). What do we call a third variable, like heredity in this case that is related to both the explanatory and response variable? In what way does the knowledge of this possible genetic link affect how we interpret the data from the study?  Assign 1, 2, 3 or 4 to the variable `heredity` in the code cell below.

1. The third variable is called an independent variable. It has no effect on the interpretation of the study’s results.
2. The third variable is called a confounding variable. It suggests that the observed association between night light use and myopia might be due to heredity rather than the light exposure itself. 
3. The third variable is called a response variable. It confirms that night light exposure is the primary cause of myopia.
4. The third variable is called a lurking variable. It strengthens the claim that night lights cause myopia.


```python
heredity = ...
```

```python
grader.check("q5_4")
```

## 6. Studying the Survivors

The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong.  (The history of the event is detailed [here](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? Assign either 1, 2, or 3 to the name `survivor_answer` below. 

1. If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2. Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.

3. Through considering the survivors, Whitehead could have identified a cure for cholera.


```python
survivor_answer = ...
```

```python
grader.check("q6_1")
```

**Note:** Whitehead ended up finding further proof that the Broad Street pump played a central role in spreading the disease to the people who lived near it. Eventually, he became one of Snow’s greatest defenders.

## Congratulations! You're done with Homework 1! 

Be sure to run all of the cells and the grader checks and verify that they all pass, then choose **Download as PDF via LaTeX** from the **File** menu, as well as **Download as Notebook** from the **File** menu, correctly name your files, and submit the two files on **canvas**.


Tests for q1 (points: [0])
```
>>> 1 <= adt_ad <= 3
True
```


Tests for q2_1 (points: [0])
```
>>> 1 <= characters_q1 <= 5
True
```


Tests for q2_2 (points: [0])
```
>>> 1 <= characters_q2 <= 3
True
```


Tests for q3_1 (points: [0])
```
>>> 1 <= names_q1 <= 4
True
```


Tests for q3_2 (points: [0])
```
>>> 1 <= names_q2 <= 4
True
```


Tests for q3_3 (points: [0])
```
>>> 1 <= names_q3 <= 4
True
```


Tests for q4_1 (points: [0])
```
>>> isinstance(biggest_change, (int, float))
True
```


Tests for q4_2 (points: [0, 0])
```
>>> isinstance(smallest_change_major, (int, float))
True
```
```
>>> 1 <= smallest_change_major <= 3
True
```


Tests for q4_3 (points: [0, 0, 0, 1, 1, 1])
```
>>> isinstance(gws_relative_change, (int, float))
True
```
```
>>> isinstance(linguistics_relative_change, (int, float))
True
```
```
>>> isinstance(rhetoric_relative_change, (int, float))
True
```
```
>>> gws_relative_change >= 64 and gws_relative_change <= 66
True
```
```
>>> linguistics_relative_change >= 36 and linguistics_relative_change <= 38
True
```
```
>>> rhetoric_relative_change >= 50 and rhetoric_relative_change <= 52
True
```


Tests for q4_4 (points: [0])
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```


Tests for q5_1 (points: [0])
```
>>> 1 <= study_type <= 2
True
```


Tests for q5_2 (points: [0])
```
>>> 1 <= associated <= 4
True
```


Tests for q5_3 (points: [0])
```
>>> 1 <= causal_relationship <= 3
True
```


Tests for q5_4 (points: [0])
```
>>> 1 <= heredity <= 4
True
```


Tests for q6_1 (points: [0])
```
>>> 1 <= survivor_answer <= 3
True
```
